# Sprint Challenge

After a sprint of Natural Language Processing, you've learned some cool new stuff: how to process text, how turn text into vectors, and how to model topics from documents. Apply your newly acquired skills to one of the most famous NLP datasets out there: [Yelp](https://www.yelp.com/dataset). As part of the job selection process, some of my friends have been asked to create analysis of this dataset, so I want to empower you to have a head start.  

The real dataset is massive (almost 8 gigs uncompressed). The data is sampled for you to something more manageable for the Sprint Challenge. You can analyze the full dataset as a stretch goal or after the sprint challenge.

## Challenge Objectives
Successfully complete all these objectives to earn full credit.

**Successful completion is defined as passing all the unit tests in each objective.**  

There are 8 total possible points in this sprint challenge.


There are more details on each objective further down in the notebook.*
* <a href="#p1">Part 1</a>: Write a function to tokenize the yelp reviews
* <a href="#p2">Part 2</a>: Create a vector representation of those tokens
* <a href="#p3">Part 3</a>: Use your tokens in a classification model on Yelp rating
* <a href="#p4">Part 4</a>: Estimate & Interpret a topic model of the Yelp reviews

____

# Before you submit your notebook you must first

1) Restart your notebook's Kernel

2) Run all cells sequentially, from top to bottom, so that cell numbers are sequential numbers (i.e. 1,2,3,4,5...)
- Easiest way to do this is to click on the **Cell** tab at the top of your notebook and select **Run All** from the drop down menu.

3) **Comment out the cell that generates a pyLDAvis visual in objective 4 (see instructions in that section).**
____

### Part 0: Import Necessary Packages
For this section, you will need to import:
- `spacy`
- `Pandas`
- `Seaborn`
- `Matplotlib`
- `NearestNeighbors`
- `Pipeline`
- `TfidfVectorizer`
- `KneighborsClassifier`
- `GridSearchCV`
- `corpora`
- `LdaModel`
- `gensim`
- `re`

> **Note: This assignment is optimized to work with these specific packages. You can use import different packages, but note that this may affect how CodeGrade works, and may cause CodeGrade to fail.**

<!-- grade_id: cell-b29df5c5bfb8c0d8, points: 0, grade: false, solution: true -->
```python
import pandas as pd

```

```python
import spacy
nlp = spacy.load('en_core_web_sm')
```

<!-- grade_id: cell-be1ef923d085ceb5, points: 1, grade: true, solution: false -->
```python
# Visible Testing
assert pd.__package__ == 'pandas'

```



### Part 0: Import Data

<!-- grade_id: cell-395851cd95d17235, points: 0, grade: false, solution: true -->
```python


# Load reviews from URL
data_url = 'https://raw.githubusercontent.com/bloominstituteoftechnology/data-science-practice-datasets/main/unit_4/unit1_nlp/review_sample.json'

# Import data into a DataFrame named df
# YOUR CODE HERE
df = pd.read_json(data_url,lines = True)
```

```python
df.head()
```

Output:
```
business_id  cool                date  funny  \
0  nDuEqIyRc8YKS1q1fX0CZg     1 2015-03-31 16:50:30      0   
1  eMYeEapscbKNqUDCx705hg     0 2015-12-16 05:31:03      0   
2  6Q7-wkCPc1KF75jZLOTcMw     1 2010-06-20 19:14:48      1   
3  k3zrItO4l9hwfLRwHBDc9w     3 2010-07-13 00:33:45      4   
4  6hpfRwGlOzbNv7k5eP9rsQ     1 2018-06-30 02:30:01      0   

                review_id  stars  \
0  eZs2tpEJtXPwawvHnHZIgQ      1   
1  DoQDWJsNbU0KL1O29l_Xug      4   
2  DDOdGU7zh56yQHmUnL1idQ      3   
3  LfTMUWnfGFMOfOIyJcwLVA      1   
4  zJSUdI7bJ8PNJAg4lnl_Gg      4   

                                                text  useful  \
0  BEWARE!!! FAKE, FAKE, FAKE....We also own a sm...      10   
1  Came here for lunch Togo. Service was quick. S...       0   
2  I've been to Vegas dozens of times and had nev...       2   
3  We went here on a night where they closed off ...       5   
4  3.5 to 4 stars\n\nNot bad for the price, $12.9...       5   

                  user_id  
0  n1LM36qN
```

```python
df.columns
```

Output:
```
Index(['business_id', 'cool', 'date', 'funny', 'review_id', 'stars', 'text',
       'useful', 'user_id'],
      dtype='object')
```

```python
df.shape
```

Output:
```
(10000, 9)
```

```python
df.cool.value_counts()
```

Output:
```
0     7513
1     1524
2      442
3      189
4      104
5       56
6       51
7       20
9       16
8       15
10      14
15       9
11       8
12       6
13       4
18       3
16       2
23       2
21       2
17       2
20       2
14       2
19       2
50       1
53       1
36       1
54       1
27       1
32       1
33       1
61       1
40       1
24       1
26       1
30       1
Name: cool, dtype: int64
```

```python
df.funny.value_counts()
```

Output:
```
0     7927
1     1216
2      416
3      168
4       79
5       54
6       30
7       28
8       17
9       13
10       8
11       8
12       8
17       4
23       3
14       2
15       2
13       2
16       2
79       1
27       1
22       1
37       1
20       1
19       1
35       1
18       1
31       1
21       1
38       1
32       1
26       1
Name: funny, dtype: int64
```

```python
df.stars.value_counts()
```

Output:
```
5    4462
4    2185
1    1496
3    1098
2     759
Name: stars, dtype: int64
```

```python
df.useful.value_counts()
```

Output:
```
0     5259
1     2197
2     1018
3      524
4      304
5      197
6      119
7       78
8       69
9       41
10      39
13      22
14      17
11      15
12      12
15      11
18       8
16       8
20       6
22       5
19       5
33       5
21       4
17       4
26       4
24       4
36       3
32       2
27       2
47       2
41       1
46       1
62       1
25       1
39       1
38       1
29       1
30       1
23       1
31       1
50       1
53       1
51       1
37       1
28       1
44       1
Name: useful, dtype: int64
```

```python
df.text.head(10)
```

Output:
```
0    BEWARE!!! FAKE, FAKE, FAKE....We also own a sm...
1    Came here for lunch Togo. Service was quick. S...
2    I've been to Vegas dozens of times and had nev...
3    We went here on a night where they closed off ...
4    3.5 to 4 stars\n\nNot bad for the price, $12.9...
5    Tasty, fast casual Latin street food.  The men...
6    This show is absolutely amazing!! What an incr...
7    Came for the Pho and really enjoyed it!  We go...
8    Absolutely the most Unique experience in a nai...
9    Wow. I walked in and sat at the bar for 10 min...
Name: text, dtype: object
```

<!-- grade_id: cell-cb5006475e42b8f9, points: 1, grade: true, solution: false -->
```python
# Visible Testing
assert isinstance(df, pd.DataFrame), 'df is not a DataFrame. Did you import the data into df?'
assert df.shape[0] == 10000, 'DataFrame df has the wrong number of rows.'
```

## Part 1: Tokenize Function
<a id="#p1"></a>

Complete the function `tokenize`. Your function should
- Accept one document at a time
- Return a list of tokens

You are free to use any method you have learned this week.

**TO PASS CODEGRADE RUNTIME:**
- Do not run your tokenize function more than one time in your notebook! It is not needed until Part 4!

<!-- grade_id: cell-3d570d5a1cd6cb64, points: 0, grade: false, solution: true -->
```python
def tokenize(doc):
    tokens = []
    for token in nlp(doc):
        if (token.is_stop == False) & (token.is_punct == False):
            tokens.append(token.lemma_)
    return tokens
```

<!-- grade_id: cell-02da164f6fbe730a, points: 1, grade: true, solution: false -->
```python
'''Testing'''
assert isinstance(tokenize(df.sample(n=1)["text"].iloc[0]), list), "Make sure your tokenizer function accepts a single document and returns a list of tokens!"
```

## Part 2: Vector Representation
<a id="#p2"></a>
1. Create a vector representation of the reviews (i.e. create a doc-term matrix).
    * Name that doc-term matrix `dtm`

<!-- grade_id: cell-0e96491cb529202c, points: 0, grade: false, solution: true -->
```python
%%time
# YOUR CODE HERE
from sklearn.feature_extraction.text import TfidfVectorizer

tfid_vect = TfidfVectorizer(stop_words="english", max_features=1000, ngram_range=(1,2))

dtm = tfid_vect.fit_transform(df['text'])

dtm = pd.DataFrame(data=dtm.toarray(), columns=tfid_vect.get_feature_names_out())
```

Output:
```
CPU times: user 5.16 s, sys: 265 ms, total: 5.43 s
Wall time: 8.75 s
```

```python
dtm.head()
```

Output:
```
00   10  10 minutes  100   11        12        15  15 minutes        20  \
0  0.0  0.0         0.0  0.0  0.0  0.000000  0.000000         0.0  0.166002   
1  0.0  0.0         0.0  0.0  0.0  0.000000  0.000000         0.0  0.000000   
2  0.0  0.0         0.0  0.0  0.0  0.000000  0.000000         0.0  0.000000   
3  0.0  0.0         0.0  0.0  0.0  0.000000  0.000000         0.0  0.000000   
4  0.0  0.0         0.0  0.0  0.0  0.170251  0.152531         0.0  0.000000   

   20 minutes  ...  write  wrong      year  years  years ago  yelp  yes  \
0    0.211222  ...    0.0    0.0  0.000000    0.0        0.0   0.0  0.0   
1    0.000000  ...    0.0    0.0  0.000000    0.0        0.0   0.0  0.0   
2    0.000000  ...    0.0    0.0  0.122783    0.0        0.0   0.0  0.0   
3    0.000000  ...    0.0    0.0  0.000000    0.0        0.0   0.0  0.0   
4    0.000000  ...    0.0    0.0  0.000000    0.0        0.0   0.0  0.0   

      young  yum  yummy  
0  0.000000  0.0    0.0  
1  0.000000  0.0    0.0  

```


2. Write a fake review. Assign the text of the review to an object called `fake_review`.
3. Query the fake review for the 10 most similar reviews, print the text of the reviews.
    - Given the size of the dataset, use `NearestNeighbors` model for this. Name the model `nn`.

<!-- grade_id: cell-3d5bc610a8ec6b24, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# Create and fit a NearestNeighbors model named "nn"
from sklearn.neighbors import NearestNeighbors

# Fit on DTM
nn = NearestNeighbors(n_neighbors=10, algorithm='kd_tree')
nn.fit(dtm)
```

Output:
```
NearestNeighbors(algorithm='kd_tree', n_neighbors=10)
```

<!-- grade_id: cell-c43704dcff67e99b, points: 1, grade: true, solution: false -->
```python
'''Testing.'''
assert nn.__module__ == 'sklearn.neighbors._unsupervised', ' nn is not a NearestNeighbors instance.'
assert nn.n_neighbors == 10, 'nn has the wrong value for n_neighbors'
```

<!-- grade_id: cell-496203e8746296ca, points: 0, grade: false, solution: true -->
```python
# Create a fake review and find the 10 most similar reviews

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-203092260fb65165, points: 1, grade: true, solution: false -->
```python
# Visible Testing
assert isinstance(fake_review, str), "Did you write a review in the correct data type?"
```

## Part 3: Classification
<a id="#p3"></a>
Your goal in this section will be to predict `stars` from the review dataset.

1. Create a pipeline object with a sklearn `CountVectorizer` or `TfidfVector` and any sklearn classifier.
    - Use that pipeline to train a model to predict the `stars` feature (i.e. the labels).
    - Use that pipeline to predict a star rating for your fake review from Part 2.



2. Create a parameter dict including `one parameter for the vectorizer` and `one parameter for the model`.
    - Include 2 possible values for each parameter
        - **Keep the values for each parameter low. Extreme values will compromise runtime**
    - **Use `n_jobs` = 1**
    - Due to limited computational resources on CodeGrader `DO NOT INCLUDE ADDITIONAL PARAMETERS OR VALUES PLEASE.`
    
    
3. Train the entire pipeline with a GridSearch
    - Name your GridSearch object as `gs`

<!-- grade_id: cell-e2beb0252d274bba, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d07134c6fe5d056e, points: 1, grade: true, solution: false -->
```python
# Visible Testing
prediction = gs.predict(["This is your prediction statement."])[0]
assert prediction in df.stars.values, 'You gs object should be able to accept raw text within a list. Did you include a vectorizer in your pipeline?'
```

## Part 4: Topic Modeling

Let's find out what those yelp reviews are saying! :D

1. Estimate a LDA topic model of the review text
    - Set num_topics to `5`
    - Name your LDA model `lda`
2. Create 1-2 visualizations of the results
    - You can use the most important 3 words of a topic in relevant visualizations.
3. In markdown, write 1-2 paragraphs of analysis on the results of your topic model

When you instantiate your LDA model, it should look like this:

```python
lda = LdaModel(corpus=corpus,
               id2word=id2word,
               random_state=723812,
               num_topics = num_topics,
               passes=1
              )

```

__*Note*__: You can pass the DataFrame column of text reviews to gensim. You do not have to use a generator.

## Note about  pyLDAvis

**pyLDAvis** is the Topic modeling package that we used in class to visualize the topics that LDA generates for us.

You are welcomed to use pyLDAvis if you'd like for your visualization. However, **you MUST comment out the code that imports the package and the cell that generates the visualization before you submit your notebook to CodeGrade.**

Although you should leave the print out of the visualization for graders to see (i.e. comment out the cell after you run it to create the viz).

### 1. Estimate a LDA topic model of the review text

* Use the `tokenize` function you created earlier to create tokens.
* Create an `id2word` object.
> Hint: Use `corpora.Dictionary`
* Create a `corpus` object.
> Hint: Use `id2word.doc2bow`
* Instantiate an `lda` model.

>> Remember to read the LDA docs for more information on the various class attributes and methods available to you in the LDA model: https://radimrehurek.com/gensim/models/ldamodel.html

```python
# Do not change this value
num_topics = 5
```

<!-- grade_id: cell-66331a185ff52f15, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

#### Testing

<!-- grade_id: cell-5a3c181311134fa9, points: 1, grade: true, solution: false -->
```python
# Visible Testing

assert lda.get_topics().shape[0] == 5, 'Did your model complete its training? Did you set num_topics to 5?'
```

#### 2. Create 2 visualizations of the results:
1. Create a visualization using pyLDAvis. Run the cell, then comment out your code before submission, leaving the visualization in the cell.

2. Create a visualization using the matplotlib library and utilizing the subplots function. Assign this visualization to a variable called `visual_plot`.


```python
# Cell for pyLDAvis visualization
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9b043e992fbd218c, points: 0, grade: false, solution: true -->
```python
# Cell for matplotlib visualzation
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f5fa579a25122b47, points: 1, grade: true, solution: false -->
```python
# Visible Testing
assert visual_plot.__module__ == 'matplotlib.axes._subplots', "You must create and assign to visual_plot a visualization"

```